<a href="https://colab.research.google.com/github/gt-cse-6040/coffee-poc/blob/main/main.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

```python
!pip install -q git+https://github.com/albert-r-waldron/cse6040-devkit.git@v1.0.0
!pip install -q dill
!wget https://github.com/gt-cse-6040/coffee-poc/archive/refs/heads/main.zip
!unzip -q main.zip
!mv coffee-poc-main/* .
!rm main.zip coffee-poc-main -r
```

Output:
```
Installing build dependencies ... [?25l[?25hdone
  Getting requirements to build wheel ... [?25l[?25hdone
  Preparing metadata (pyproject.toml) ... [?25l[?25hdone
  Building wheel for cse6040_devkit (pyproject.toml) ... [?25l[?25hdone
[2K   [90m━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━[0m [32m119.4/119.4 kB[0m [31m2.5 MB/s[0m eta [36m0:00:00[0m
[?25h--2025-02-09 20:06:54--  https://github.com/gt-cse-6040/coffee-poc/archive/refs/heads/main.zip
Resolving github.com (github.com)... 140.82.114.4
Connecting to github.com (github.com)|140.82.114.4|:443... connected.
HTTP request sent, awaiting response... 302 Found
Location: https://codeload.github.com/gt-cse-6040/coffee-poc/zip/refs/heads/main [following]
--2025-02-09 20:06:54--  https://codeload.github.com/gt-cse-6040/coffee-poc/zip/refs/heads/main
Resolving codeload.github.com (codeload.github.com)... 140.82.113.10
Connecting to codeload.github.com (codeload.github.com)|140.82.113.10|:443... connected.
HTTP request sent, 
```

# `assignment name`: `assignment subtitle`
_Version 0.0.1_

*All of the header information is important. Please read it..*

**Topics number of exercises:** This problem builds on your knowledge of `this, that, and the other`. It has **1** exercises numbered 0 to **0**. There are **total points** available points. However to earn 100% the threshold is **points cap** points. (Therefore once you hit **points cap** points you can stop. There is no extra credit for exceeding this threshold.)

**Exercise ordering:** Each exercise builds logically on previous exercises but you may solve them in any order. That is if you can't solve an exercise you can still move on and try the next one. Use this to your advantage as the exercises are **not** necessarily ordered in terms of difficulty. Higher point values generally indicate more difficult exercises.

**Demo cells:** Code cells starting with the comment `### Run Me!!!` load results from prior exercises applied to the entire data set and use those to build demo inputs. These must be run for subsequent demos to work properly but they do not affect the test cells. The data loaded in these cells may be rather large (at least in terms of human readability). You are free to print or otherwise use Python to explore them but we may not print them in the starter code.

**Debugging your code:** Right before each exercise test cell there is a block of text explaining the variables available to you for debugging. You may use these to test your code and can print/display them as needed (careful when printing large objects you may want to print the head or chunks of rows at a time).

**Exercise point breakdown:**


- Exercise 0 - : **1** point(s)


**Final reminders:**

- Submit after **every exercise**
- Review the generated grade report after you submit to see what errors were returned
- Stay calm, skip problems as needed and take short breaks at your leisure

```python
### Global imports
import dill
from cse6040_devkit import plugins, utils
import re
import pandas as pd


```

### Exercise 0: (1 points)
**reverse_dict**  

**Your task:** define `reverse_dict` as follows:



```python
### Solution - Exercise 0
def reverse_dict(d):
    ### BEGIN SOLUTION
    rev_d = {v:k for k, v in d.items()}
    if len(d) != len(rev_d):
        raise ValueError
    return rev_d
    ### END SOLUTION

### Demo function call
valid_d = {'A': 1, 'B': 2, 'C': 3}
invalid_d = {'A': 1, 'B': 2, 'C': 2}
error_raised, rev_d = plugins.error_handler(reverse_dict)(valid_d)
print(f'Results of `reverse_dict`(valid_d): Error Raised {error_raised}; Result {rev_d}')
error_raised, rev_d = plugins.error_handler(reverse_dict)(invalid_d)
print(f'Results of `reverse_dict`(invalid_d): Error Raised {error_raised}; Result {rev_d}')
```

Output:
```
Results of `reverse_dict`(valid_d): Error Raised False; Result {1: 'A', 2: 'B', 3: 'C'}
Results of `reverse_dict`(invalid_d): Error Raised True; Result None
```



**The demo should display this printed output.**
```
Results of `reverse_dict`(valid_d): Error Raised False; Result {1: 'A', 2: 'B', 3: 'C'}
Results of `reverse_dict`(invalid_d): Error Raised True; Result None
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for reverse_dict (exercise 0). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output.


<!-- grade_id: ex_0, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 0

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    execute_tests = dill.load(f)

# Execute test
passed, test_case_vars, e = execute_tests(func=plugins.error_handler(reverse_dict),
              ex_name='reverse_dict',
              key=b'QjweeQJQjJGkvXjTnF1i7M5ibvMV_Xtqv4Q2py5RR8w=',
              n_iter=100)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to reverse_dict did not pass the test.'

### BEGIN HIDDEN TESTS
passed, test_case_vars, e = execute_tests(func=plugins.error_handler(reverse_dict),
              ex_name='reverse_dict',
              key=b'NDTF82LIgRJW82C7H86Rblz4dx3pv19JrQTpGsGzRFU=',
              n_iter=1,
              hidden=True)
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
if e: raise e
assert passed, 'The solution to reverse_dict did not pass the test.'
### END HIDDEN TESTS
print('Passed! Please submit.')
```

Output:
```
reverse_dict test ran 99 iterations in 0.04 seconds
reverse_dict test ran 0 iterations in 0.01 seconds
Passed! Please submit.
```